---
title: 'Loading the Data Set'
description: 'Loading Wisconsin Breast Cancer Data Set'
---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: a661985a47
xp: 50
```



`@possible_answers`
1. Yes
2. No
[3. Sometimes]

`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```
